Before you turn this problem in, make sure everything runs as expected. First, **restart the kernel** (in the menubar, select Kernel$\rightarrow$Restart) and then **run all cells** (in the menubar, select Cell$\rightarrow$Run All).

Make sure you fill in any place that says `YOUR CODE HERE` or "YOUR ANSWER HERE", as well as your name and collaborators below:

```python
NAME = "Jingren Wang"
COLLABORATORS = "N.A."
```

---

# CS110 Pre-class Work 1.2



## Question 1
The following pseudocode has been extracted from Cormen et al. You should analyze the lines of the pseudocode trying to understand what every line in the pseudocode does and how many times is that line executed. 
![Insertion sort pseudo-code](insertionsort_pseudo.png)

Answer the following questions: What does line 7 do? Why is it necessary?

<!-- grade_id: cell-3bac59ff24b11a27, points: 0, grade: true, solution: true -->
1. the purpose of line 5-7 is to continue comparing the number at hand (key) with the biggest number in the rest of uncompared ones in the sorted list A[1...i-1], if key is smaller, then execute line 6 to move the current biggest number forward 1 index, all the way until the current biggest is smaller than the key. So, line 7 updates the index 1 step left of the sorted list in order to compare the next biggest number with the key. line 7 is necessary, a removal of which would render an infinite while loop.
2. in an extreme case where the key is smaller than any number in the sorted list, the value of i in last iterated loop would be i=1 such that the A[1] item is moved to A[2] position, then i = 1-1 =, 0, which renders i>0 false, while loop exited and key value is assigned to position A[0+1] = A[1]
* be minded that in implementation with python, indexing starts from 0

## Question 2
The following Python code is an attempt to implement the Insertion-Sort pseudocode above. However, this code has a bug. 


```python
def insertionSort(A):
    for j in range(len(A)):
        key = A[j]
        i= j-1
        while i >= 0 and A[i]>key:
            A[i+1] = A[i]
            i -= 1
            A[i+1] = key
    return A
```

Copy and paste the code in the cell below. Identify and correct the bug. 

<!-- grade_id: cell-1b7488e89cb501c5, points: 0, grade: false, solution: true -->
```python
def insertionSort(A):
    # introduce global counter variable
    global count
    count = 0
    
    for j in range(1,len(A)):
        key = A[j]
        count += 1
        i= j-1
        count += 1
        while i >= 0 and A[i]>key:
            count += 1  # count step for while statement evaluation
            A[i+1] = A[i]
            count += 1
            i -= 1
            count += 1
        A[i+1] = key
        count += 1
        
    return A
    raise NotImplementedError()
    
count = 0
print(insertionSort([0]) == [0])
print(insertionSort([-1,1]) == [-1,1])
print(insertionSort([1,-1]) == [-1,1])
print(insertionSort([1,6,3,6]) == [1,3,6,6])

print(count)


```

Output:
```
True
True
True
True
12
```

<!-- grade_id: cell-3432ce8ae2493fa6, points: 1, grade: true, solution: false -->
```python
assert(insertionSort([0]) == [0])
assert(insertionSort([-1,1]) == [-1,1])
assert(insertionSort([1,-1]) == [-1,1])
assert(insertionSort([1,6,3,6]) == [1,3,6,6])
```

How do you know this code works as intended now?

<!-- grade_id: cell-442998d5fdd8b561, points: 0, grade: true, solution: true -->
1. Bug: line 8 was not corrected indented, final assignment of key (in hand) should be executed outside of while loop, when correct i is found such that A[i]< key.
2. To check if the debugged code runs well, we could come up with some test cases of A, and use assert()to check boolean status, if false, an assertion error would raised.
3. To further prove that the sorting outcome is correct, we neet to check:
a. the sorted A is of same length as the unsorted one;
b. A[n-1]< A[n] for n in {1,...A.len}

```python
import time

def get_InsertionSort_data():
    n_list = []  # list of input size, n = 100*k
    steps_list = [] # list of step for each k
    runTime_list = [] # list of running time for each k
    global count
    
    for k in range(1,15): 
        count = 0
        
        list_k = [i for i in range(100*k, 0, -1)]
        p = 1
        r = len(list_k)
        n_list.append(r)

        start_time = time.clock()
        insertionSort(list_k)
        runTime = time.clock()-start_time
        runTime_list.append(runTime)

        steps_list.append(count)
        
    return n_list,steps_list,runTime_list

    raise NotImplementedError()

```

```python
def get_Insertion_Steps(A):
    
    insertionSort(A)
    step = count
    
    return step

A = [6,5,3,2]
stp = get_Insertion_Steps(A)
print(stp)
```

Output:
```
27
```

## Question 3 
Implement the bubble sort algorithm in the cell below. First, look up the corresponding pseudocode in Cormen et al. (page 40), and write a Python implementation for the bubble sort.  Like in the `insertionSort`,  your new Python function, `bubbleSort`, takes a list of elements as an argument and returns the sorted list.

<!-- grade_id: cell-b709733c96d8b615, points: 0, grade: false, solution: true -->
```python
def bubbleSort(A):
    global count
    b = 0 # initialize a bin variable for temporary keeping of swapped item
    count += 1
    n = len(A) #length of list items
    count += 1
    for i in range(n-1):  # i from 0 to n-1
        count += 1
        for j in range(n-1,i,-1): # j from n-1 to 1(for first round)
            count += 1
            if A[j]< A[j-1]:
                count += 1
                b = A[j-1]
                count += 1
                A[j-1] = A[j]
                count += 1
                A[j] = b   # three-step swap
                count += 1
    return A       
    raise NotImplementedError()
```

```python
A = [5,2,6,3,1,0] 
bubbleSort(A) 
print(count)
print(A)

```

Output:
```
91
[0, 1, 2, 3, 5, 6]
```

<!-- grade_id: cell-ffc9b251ea0047ba, points: 2, grade: true, solution: false -->
```python
# Please ignore this cell. This cell is for us to implement the tests 
# to see if your code works properly. 
```

Is your program bug-free? Think of three different [“test cases”](https://www.hiredintech.com/classrooms/algorithm-design/lesson/83) to check the correctness of your code. If a code passes the tests you have just outlined, is this an unambiguous proof that the code is correct? Give your answer in the cell below. Feel free to use additional cells (Markdown or Code) as you wish.

<!-- grade_id: cell-5ecb51404140cb2f, points: 0, grade: true, solution: true -->
No, even if my code passes all three types of test cases, it is not a guaranteed proof since there will be cases left out from sample cases, or even extensive cases. Also, edge case testing and load testing may vary from time to time as the computer hardware wears out.

```python
## sample test
assert(bubbleSort([0]) == [0])
assert(bubbleSort([-1,1]) == [-1,1])
assert(bubbleSort([1,-1]) == [-1,1])
assert(bubbleSort([1,6,3,6]) == [1,3,6,6])
assert(bubbleSort([5,2,6,3,1,0]) == [0, 1, 2, 3, 5, 6])
# worst case:
assert(bubbleSort([9,8,7,6,5,4,3,2,1,0]) == [0,1,2,3,4,5,6,7,8,9])

### no assertion error raised!
```

```python
## extensive test - non-trivial functional test
# logic: A is well sorted if A[i]<= A[i+1] for i in 0 to n-1
def bubbleCheck(A, len_A):
    if len(A) != len_A:
        return "length unequal."
    else: 
        for i in range(len_A-1):
            if A[i]> A[i+1]:
                return "not well sorted."
    return "well sorted!"

A = [9,8,7,6,5,4,3,2,1,0]
len_A = len(A)
bubbleSort(A)
bubbleCheck(A, len_A)

```

## Question 4
The following pseudocode corresponds to a third sorting algorithm to be covered in Lesson 1.2, the selection sort.  Like in the previous step, write a Python implementation of the algorithm and make sure that your program is bug-free. `selectionSort` should take a list of elements as an argument and return the sorted list.

![Selection sort pseudo-code](selectionsort_pseudo.png)

<!-- grade_id: cell-e04f94e5d418923c, points: 0, grade: false, solution: true -->
```python
def selectionSort(A):
    b = 0 # initialize a b as temporary item holder during swapping
    n = len(A)
    for i in range(n-1):  # i in 0 to n-1
        min_idx = i  # index of default minimal 
        for j in range(i+1,n):   # j in i+1 to n
            if A[j] < A[min_idx]:
                min_idx = j     #update minimal index
            # swap A[i] with A[min_idx]  
            b = A[min_idx]
            A[min_idx] = A[i]
            A[i] = b   # three-step swap
    return A
    raise NotImplementedError()
```

```python
## extensive test - non-trivial functional test
# logic: A is well sorted if A[i]<= A[i+1] for i in 0 to n-1
def selectionCheck(A, len_A):
    if len(A) != len_A:
        return "length unequal."
    else: 
        for i in range(len(A)-1):
            if A[i]> A[i+1]:
                return "not well sorted."
    return "well sorted!"

A = [9,8,7,6,5,4,3,2,1,0]
len_A = len(A)
selectionSort(A)
selectionCheck(A, len_A)

```

```python

```

<!-- grade_id: cell-ffe912ac7922de70, points: 1, grade: true, solution: false -->
```python
# Please ignore this cell. This cell is for us to implement the tests 
# to see if your code works properly. 
```

**Please make sure you try your best to code insertion sort, selection sort and bubble sort, and bring these programs to class. You will need them for a class activity.**