Before you turn this problem in, make sure everything runs as expected. First, **restart the kernel** (in the menubar, select Kernel$\rightarrow$Restart) and then **run all cells** (in the menubar, select Cell$\rightarrow$Run All).

Make sure you fill in any place that says `YOUR CODE HERE` or "YOUR ANSWER HERE", as well as your name and collaborators below:

```python
NAME = "Jingren Wang"
COLLABORATORS = "N.A."
```

---

# CS110 Pre-class Work 3.1

## Question 1. 
Paste in your Python implementation of the maximum subarray from the previous class in the cell below and use that to find out the value of the maximum subarray of this array: `A = [-2, -3, 4, -1, -2, 1]`

```python
import math
import numpy as np


def bruteforce_max_subarray(A):
    """Implements brute-force maximum subarray finding.
    
    Inputs:
    - A: a NON-EMPTY list of floats
    
    Outputs: A tuple of
    - the start index of the max subarray
    - the end index of the max subarray
    - the value of the maximum subarray
    """
    
    max_sum = -math.inf #give floor max (-oo) for initial max_sum
    max_i, max_j = 0,0 
        
    for i in range(len(A)-1):
        # print("i = ", i)
        for j in range(i+1, len(A)):
            # print("  j = ",j)
            subArray_sum = sum(A[i:j])
            # print(A[i:j])
            if subArray_sum >= max_sum:
                max_sum = subArray_sum
                max_i = i
                max_j = j # minus 1 to get the correct right index of current subarray.
                
    return (max_i,max_j,max_sum)

    raise NotImplementedError()
```

<!-- grade_id: cell-e69c8bbcd40ca242, points: 0, grade: true, solution: true -->
```python
# YOUR CODE HERE
A = [-2, -3, 4, -1, -2, 1]
sumMax_A = bruteforce_max_subarray(A)
print(sumMax_A)

#raise NotImplementedError()
```

Output:
```
(2, 3, 4)
```

## Question 2. 
Now, your friend Joe comes and appends a single extra number at the end of the array, which becomes:  `B = [-2, -3, 4, -1, -2, 1, 8]`. Do you have to re-run the entire maximum subarray again? Explain your answer. 
The subsequent questions will help you figure out an efficient algorithmic strategy to address the last question, but make sure to write your explanation above first, before answering the remaining questions.


<!-- grade_id: cell-d4b7cd845c816263, points: 0, grade: true, solution: true -->
My answer:<br>
Yes, the bruteforce_max_subarray() function is 'stupid' in the sense that it exhaust every possible combination of (i,j) of A[i:j] for any input array A. So the function would not be able to identify A in B (as A+'8'), and blindly treat array B as the new array A, which is quite inefficient.

## Question 3. 

**Determine if the following statement is True or False and explain your answer.**
If the maximum subarray of the array A is different than the maximum subarray of the array B (questions 1 and 2), the new maximum-subarray doesn’t need to include 8 (i.e., the newly appended element). 


<!-- grade_id: cell-9377964a8756f13b, points: 0, grade: true, solution: true -->
My answer:<br>

Given the key assumption of a maximum subarray probelm, that a subarray must contain contiguous elements of A, the question of whether the maximum subarray of array B would differ from that of array A can be reduced to a comparison of the sum of value from elements to the left of the maximum subarray with zero.
e.g. for A = [-2, -3, 4, -1, -2, 1]
the subarray that returns a maximum sum is A[2:3] = [4], such that the right index j = 3 
next, the maximum subarray value will be updated ONLY if the sum of all elements of A right to jth element, plus the added one, is greater than 0
    in this case A[j:len(A)] = [-1,-2,1], whose sum is -1-2+1 = -2,
since 8+(-2) = 6 > 0, a new sum is obtained and the maximus subarray must include 8.

Therefore, this statement is false, as the new maximu-subarray MUST include 8, the newly added element, if the new array B returns a different (must be higher) value than the current maximum of array A.

## Question 4.
Complete the Python function `incremental_max_subarray(x, mx)` in the cell below.

This [video](https://www.youtube.com/watch?v=AAgErqQmwmA&list=PLF_a-qBXTGFektoI6JUOTRL36JlvD04BR&index=4&t=0s)  might be helpful to understand the `incremental_max_subarray` problem.

<!-- grade_id: cell-0230e459f3d701f8, points: 0, grade: false, solution: true -->
```python
def incremental_max_subarray(x, mx):
    """ 
    Inputs:
    - x: a NON-EMPTY list of numbers (e.g., the array B in the first two questions above). 
        * If x has 1 element: returns the value of the element regardless of the value of mx
    - mx: the maximum subarray of x excluding its last element (i.e., compute the 
    maximum subarray of the input array x considering only its first len(x) - 1 elements)
    
    Output: The sum value (in float) of the maximum subarray of the array x.
    
    For example, using the array B in question 2, the result of incremental_max_subarray(B, 4) 
    is 10 (10 = 8 + 1 - 2 -1 + 4). 
    """
    # YOUR CODE HERE
    n = len(x)  
    if n == 1: # return the element when array is singleton 
        return x[0]
    elif n == 0:  # raise error when input array is empty
        return "array is empty."
    else:
        sumArr = 0
        new_max = -math.inf # initiate sum of new array from B[-1] at floor value
        
        for i in range(n-1,-1,-1):
           # print('i=',i)
            sumArr += x[i]
            if sumArr > new_max:
                new_max = sumArr
                max_i = i
        
    return max(float(new_max), float(mx))
    
    raise NotImplementedError()
```

<!-- grade_id: cell-9abf19e56620ffa3, points: 1, grade: true, solution: false -->
```python
B = [-2, -3, 4, -1, -2, 1, 8]
assert(incremental_max_subarray(B, 4) == 10)
assert(incremental_max_subarray(B[:1], 0) == B[0])
```

<!-- grade_id: cell-bacebe71191cbd0f, points: 0, grade: true, solution: true -->
YOUR ANSWER HERE

<!-- grade_id: cell-fd96d4ccccd99fe6, points: 0, grade: false, solution: true -->
```python
def max_subarray(A):
    """
    Using `incremental_max_subarray` iteratively on A to produce the value of the maximum
    subarray of A.
    
    Inputs:
    - A: a NON-EMPTY list of floats
    
    Outputs: float, the sum of the maximum subarray of A
    """
    # YOUR CODE HERE
    mx = -math.inf  #initiate mx at floor value
    l = len(A)
    
    for j in range(1, l+1):   # upper index j from 1 to n
        x = A[:j]  # 
        mx = incremental_max_subarray(x, mx)  # assign mx to current new max 
    return mx
    
    raise NotImplementedError()

```

```python
# my test case 
n = 10 # array size
A = np.random.uniform(-20, 20, n)
np.set_printoptions(precision=2)
print('array A is ', A)

# incremental method
incr_A = max_subarray(A)
print('maximum subarray sum of A is ''%.2f' % incr_A)


```

Output:
```
array A is  [-12.16 -11.3    1.03   2.5    8.1   10.05  17.54  -9.24  -5.31  10.91]
maximum subarray sum of A is 39.22
```

```python
# [please ignore] my doodle for inspecting max_subarray(A)
def max_subarray_inspect(A):
    mx = -math.inf  #initiate mx at floor value
    l = len(A)
    
    for j in range(1, l+1):   # upper index j from 1 to n
        print ('current_mx=', mx)
        print('>>> call incremental_max_subarray(x, mx)')
        print('A=',A[:j])
        x = A[:j]  # 
        mx = incremental_max_subarray(x, mx)  # assign mx to current new max 
        print('new_mx=',mx)
        print('')
        
    return mx    

# inspection
C = [-2, -3, 4, -1, -2, 1, 4]
max_subarray_inspect(C)
```

Output:
```
current_mx= -inf
>>> call incremental_max_subarray(x, mx)
A= [-2]
new_mx= -2

current_mx= -2
>>> call incremental_max_subarray(x, mx)
A= [-2, -3]
new_mx= -2.0

current_mx= -2.0
>>> call incremental_max_subarray(x, mx)
A= [-2, -3, 4]
new_mx= 4.0

current_mx= 4.0
>>> call incremental_max_subarray(x, mx)
A= [-2, -3, 4, -1]
new_mx= 4.0

current_mx= 4.0
>>> call incremental_max_subarray(x, mx)
A= [-2, -3, 4, -1, -2]
new_mx= 4.0

current_mx= 4.0
>>> call incremental_max_subarray(x, mx)
A= [-2, -3, 4, -1, -2, 1]
new_mx= 4.0

current_mx= 4.0
>>> call incremental_max_subarray(x, mx)
A= [-2, -3, 4, -1, -2, 1, 4]
new_mx= 6.0

6.0
```

<!-- grade_id: cell-669ad779766aa2c3, points: 1, grade: true, solution: false -->
```python
# Please ignore this cell. This cell is for us to implement the tests 
# to see if your code works properly. 
```

### Part 2.
Is this more efficient than the divide-and-conquer approach? Explain.

<!-- grade_id: cell-cd8f0b130a7136db, points: 0, grade: true, solution: true -->
My answer:<br>

According to Cormen, the divide-and-conquer approach is asymptotically $\theta(nlog(n))$

We estimate the asymptotic behavior of incremental approach:
a. for an input array A of length n, the for loop of j loops n times; 
b. within each loop, the incremental_max_subarray(x,mx) function is called once, with x a subarray of A of length m = 1 to n; 
c. the incremental_max_subarray(x,mx) of input size n, for n>1, is executed n times due to the downwards counting for loop of i;
d. since the if- and assignment statements run at constant time, the dominant pattern of growth is with the for loop of j, and iterative call of incremental_max_subarray(x,mx), in which 
<br>
<br>
T(n) = (1+...+n) + big_theta(n) = n(n+1)/2 +big_theta(n) <br>
    = 0.5*n^2+ c*n
*given that, both loops iterate through every element in input array regardless of sorting, the best-case scenario would be the same in terms of order of growth, so big_theta, not big_Oh, 
the asymptotic behavior of incremental approach is $\theta(n^2)$, which is less efficient than the divide-and-conquer approach.
 




```python

```